---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Data Cleaning for Analysis in a SQL Database

```yaml
type: VideoExercise
key: 786eb72bf2
xp: 50
```

`@projector_key`
564c452a5a07e95418480fc9bc4648a9

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 1ef24b98ea
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
